<?xml version="1.0" encoding="UTF-8"?>
<quiz>
    <question type="stack">
        <name>
            <text> Uf </text>
        </name>
        <questiontext format="html">
            <text>
                <![CDATA[<p>Consider \(U_f: |x,y\rangle \to | x, y \oplus f(x)\rangle\) for a function \(f:\{0,1\} \to \{0,1\}\).</p>
<p>Give, if possible, the resulting state in ket notation. If the result depends on \(f\), write <code>depends</code>.</p>
<ul>
<li>\(U_f (| 0 \rangle \otimes |+\rangle) = \) [[input:ans1]]  [[validation:ans1]]</li>
<li>\(U_f (| 0 \rangle \otimes |-\rangle) = \) [[input:ans2]]  [[validation:ans2]]</li>
<li>\(U_f (| - \rangle \otimes |+\rangle) = \) [[input:ans3]]  [[validation:ans3]]</li>
<li>\(U_f (| 0 \rangle \otimes |0\rangle) = \) [[input:ans4]]  [[validation:ans4]]</li>
</ul>
<hr />
<p><strong>Input instructions:</strong></p>
<p>Give the answer in ket-notation.
Insert \(|000\rangle\) as <code>ket(000)</code> (that is correct, even though the preview will cut out leading 0s and show \(|0\rangle\) always),
\(|100\rangle\) as <code>ket(100)</code>, … . You can enter \(a+b\) as <code>a+b</code>, \(a*b\) as <code>a*b</code> (<code>ab</code> is not allowed),
\(\frac{a}{b}\) as <code>a/b</code>, \(a^b\) as <code>a^b</code>, \(a_b\) as <code>a_b</code>, \(\sqrt{a}\) as <code>sqrt(a)</code>, \(e\) as <code>e</code>, \(i\) as <code>i</code>.</p>
]]>
            </text>
        </questiontext>
        <generalfeedback format="html">
            <text></text>
        </generalfeedback>
        <defaultgrade>1</defaultgrade>
        <penalty>0.1</penalty>
        <hidden>0</hidden>
        <idnumber></idnumber>
        <stackversion>
            <text>2025040100</text>
        </stackversion>
        <questionvariables>
            <text>
                <![CDATA[texput(ket, lambda([ex], sconcat("| ", tex1(first(ex)), "\\rangle")))]]>
            </text>
        </questionvariables>
        <specificfeedback format="html">
            <text></text>
        </specificfeedback>
        <questionnote format="moodle_auto_format">
            <text></text>
        </questionnote>
        <questiondescription format="moodle_auto_format">
            <text></text>
        </questiondescription>
        <questionsimplify>1</questionsimplify>
        <assumepositive>0</assumepositive>
        <assumereal>0</assumereal>
        <prtcorrect format="html">
            <text> &lt;p&gt;Richtige Antwort, gut gemacht!&lt;/p&gt; </text>
        </prtcorrect>
        <prtpartiallycorrect format="html">
            <text> &lt;p&gt;Ihre Antwort ist teilweise korrekt.&lt;/p&gt; </text>
        </prtpartiallycorrect>
        <prtincorrect format="html">
            <text> &lt;p&gt;Falsche Antwort.&lt;/p&gt; </text>
        </prtincorrect>
        <decimals>.</decimals>
        <scientificnotation>*10</scientificnotation>
        <multiplicationsign>dot</multiplicationsign>
        <sqrtsign>1</sqrtsign>
        <complexno>i</complexno>
        <inversetrig>cos-1</inversetrig>
        <logicsymbol>lang</logicsymbol>
        <matrixparens>[</matrixparens>
        <variantsselectionseed></variantsselectionseed>
        <input>
            <name>ans1</name>
            <type>algebraic</type>
            <tans>1/sqrt(2)*(ket(00)+ket(01))</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>
                A, a, aA, Aa, AA, aa, ab, aB, Ab, AB, Ac, AC, aC, ac, aD, AD, Ad, ad, ae, AE, Ae, aE, Af, AF, af, aF, Ag, aG, AG, ag, aH, ah, Ah, AH, ai, aI, AI, Ai, AJ, Aj, aJ, aj, AK, Ak, aK, ak, al, aL, Al, AL, am, aM, Am, AM, AN, An, an, aN, AO, Ao, aO, ao, ap, aP, AP, Ap, aq, aQ, AQ, Aq, ar, aR, AR, Ar, aS, as, AS, As, At, AT, at, aT, Au, AU, au, aU, av, aV, Av, AV, aW, aw, Aw, AW, AX, ax, aX, Ax, AY, Ay, ay, aY, AZ, Az, aZ, az, b, B, ba, bA, BA, Ba, BB, Bb, bB, bb, bc, bC, BC, Bc, BD, Bd, bD, bd, Be, BE, bE, be, bF, bf, Bf, BF, BG, bg, bG, Bg, BH, Bh, bH, bh, bi, bI, BI, Bi, bj, bJ, Bj, BJ, Bk, BK, bk, bK, BL, Bl, bl, bL, bM, bm, BM, Bm, bN, bn, Bn, BN, Bo, BO, bO, bo, BP, Bp, bP, bp, bq, bQ, BQ, Bq, bR, br, Br, BR, bs, bS, Bs, BS, bt, bT, BT, Bt, Bu, BU, bU, bu, Bv, BV, bV, bv, bW, bw, BW, Bw, bX, bx, BX, Bx, by, bY, By, BY, Bz, BZ, bz, bZ, C, c, cA, ca, Ca, CA, cb, cB, CB, Cb, CC, Cc, cc, cC, cd, cD, CD, Cd, CE, Ce, ce, cE, CF, Cf, cF, cf, cg, cG, CG, Cg, Ch, cH, ch, CH, CI, Ci, cI, ci, cj, cJ, CJ, Cj, Ck, cK, ck, CK, Cl, CL, cL, cl, CM, cm, cM, Cm, cn, cN, Cn, CN, cO, co, Co, CO, CP, Cp, cp, cP, CQ, Cq, cq, cQ, cR, cr, CR, Cr, cS, cs, Cs, CS, ct, cT, CT, Ct, cu, cU, Cu, CU, CV, Cv, cv, cV, CW, Cw, cW, cw, cx, cX, Cx, CX, cy, cY, CY, Cy, Cz, cz, cZ, CZ, D, d, DA, Da, da, dA, dB, db, DB, Db, dC, dc, Dc, DC, Dd, DD, dD, dd, dE, de, De, DE, Df, DF, dF, df, DG, Dg, dG, dg, dH, dh, DH, Dh, DI, Di, di, dI, Dj, DJ, dj, dJ, dK, dk, Dk, DK, DL, Dl, dL, dl, Dm, DM, dm, dM, dn, dN, Dn, DN, do, dO, Do, DO, dP, DP, Dp, dp, Dq, DQ, dq, dQ, Dr, DR, dR, dr, ds, dS, Ds, DS, dt, dT, DT, Dt, dU, du, Du, DU, dV, dv, DV, Dv, DW, Dw, dW, dw, Dx, DX, dX, dx, dy, dY, Dy, DY, dz, dZ, DZ, Dz, E, Ea, EA, ea, eA, Eb, EB, eb, eB, ec, eC, EC, Ec, ed, eD, ED, Ed, Ee, EE, ee, eE, ef, eF, EF, Ef, EG, Eg, eg, eG, EH, Eh, eh, eH, ei, eI, EI, Ei, EJ, Ej, ej, eJ, ek, EK, Ek, eK, el, eL, El, EL, Em, EM, em, eM, EN, En, en, eN, Eo, eO, eo, EO, ep, eP, EP, Ep, EQ, Eq, eq, eQ, ER, Er, eR, er, ES, Es, eS, es, et, eT, Et, ET, eu, eU, Eu, EU, eV, ev, EV, Ev, ew, eW, EW, Ew, EX, Ex, ex, eX, EY, Ey, eY, ey, ez, eZ, EZ, Ez, f, F, FA, Fa, fa, fA, FB, Fb, fb, fB, Fc, FC, fC, fc, fD, fd, Fd, FD, fe, fE, FE, Fe, FF, Ff, fF, ff, fg, fG, Fg, FG, Fh, FH, fH, fh, FI, Fi, fI, fi, fJ, fj, FJ, Fj, Fk, fk, FK, fK, fL, fl, FL, Fl, fM, fm, FM, Fm, Fn, FN, fN, fn, Fo, FO, fO, fo, FP, Fp, fP, fp, fq, fQ, Fq, FQ, fr, Fr, FR, fR, Fs, FS, fs, fS, FT, Ft, ft, fT, fU, fu, Fu, FU, fV, fv, Fv, FV, fW, fw, FW, Fw, fX, fx, Fx, FX, FY, Fy, fy, fY, Fz, FZ, fz, fZ, G, g, gA, ga, Ga, GA, Gb, GB, gb, gB, Gc, GC, gc, gC, GD, Gd, gD, gd, ge, gE, Ge, GE, gF, gf, Gf, GF, Gg, GG, gG, gg, gH, gh, Gh, GH, Gi, GI, gI, gi, Gj, GJ, gJ, gj, gK, gk, Gk, GK, Gl, gl, GL, gL, gm, gM, GM, Gm, gn, gN, Gn, GN, GO, Go, gO, go, GP, Gp, gp, gP, GQ, gq, gQ, Gq, gR, gr, GR, Gr, Gs, gS, gs, GS, Gt, GT, gt, gT, GU, gu, Gu, gU, gv, gV, Gv, GV, gW, gw, GW, Gw, gx, gX, GX, Gx, gY, gy, Gy, GY, GZ, Gz, gZ, gz, H, h, ha, hA, Ha, HA, hb, hB, HB, Hb, Hc, HC, hC, hc, Hd, HD, hD, hd, HE, he, hE, He, hf, hF, HF, Hf, hG, Hg, hg, HG, Hh, HH, hh, hH, hI, Hi, hi, HI, HJ, Hj, hj, hJ, Hk, HK, hk, hK, hl, hL, HL, Hl, Hm, HM, hM, hm, hN, hn, HN, Hn, ho, hO, HO, Ho, Hp, HP, hP, hp, Hq, HQ, hQ, hq, hr, hR, Hr, HR, hs, hS, Hs, HS, HT, Ht, hT, ht, Hu, HU, hU, hu, hv, Hv, hV, HV, hW, hw, HW, Hw, hX, hx, HX, Hx, hy, hY, Hy, HY, hZ, HZ, Hz, hz, I, Ia, IA, ia, iA, ib, iB, Ib, IB, iC, ic, IC, Ic, ID, Id, id, iD, Ie, IE, ie, iE, if, If, iF, IF, iG, ig, IG, Ig, IH, Ih, iH, ih, II, Ii, iI, ii, ij, Ij, IJ, iJ, Ik, IK, iK, ik, IL, Il, il, iL, iM, im, IM, Im, IN, in, In, iN, iO, io, IO, Io, ip, iP, IP, Ip, IQ, Iq, iQ, iq, IR, Ir, ir, iR, iS, is, IS, Is, it, iT, It, IT, Iu, IU, iu, iU, Iv, IV, iv, iV, iW, iw, Iw, IW, ix, iX, Ix, IX, iy, iY, IY, Iy, iZ, iz, Iz, IZ, j, J, JA, Ja, ja, jA, JB, Jb, jB, jb, jC, jc, Jc, JC, jd, jD, Jd, JD, JE, Je, je, jE, JF, Jf, jF, jf, jG, jg, JG, Jg, jh, jH, Jh, JH, JI, ji, jI, Ji, Jj, JJ, jJ, jj, jk, jK, Jk, JK, JL, Jl, jL, jl, Jm, JM, jM, jm, jN, jn, Jn, JN, jo, jO, Jo, JO, jP, jp, Jp, JP, jq, jQ, JQ, Jq, Jr, JR, jr, jR, Js, JS, js, jS, jt, jT, JT, Jt, jU, ju, Ju, JU, Jv, JV, jV, jv, JW, Jw, jW, jw, jX, jx, JX, Jx, jy, jY, JY, Jy, jz, jZ, Jz, JZ, K, k, kA, ka, Ka, KA, Kb, KB, kb, kB, Kc, KC, kc, kC, kd, kD, KD, Kd, ke, kE, Ke, KE, Kf, KF, kF, kf, Kg, KG, kG, kg, kH, kh, KH, Kh, kI, ki, KI, Ki, kj, kJ, Kj, KJ, Kk, KK, kk, kK, kl, kL, Kl, KL, Km, KM, km, kM, KN, kn, kN, Kn, ko, kO, KO, Ko, kP, kp, KP, Kp, kQ, KQ, kq, Kq, kR, kr, Kr, KR, Ks, KS, kS, ks, Kt, KT, kt, kT, ku, kU, Ku, KU, kv, kV, KV, Kv, KW, Kw, kW, kw, Kx, KX, kx, kX, ky, kY, KY, Ky, kz, kZ, KZ, Kz, L, l, la, lA, LA, La, lb, lB, LB, Lb, Lc, LC, lC, lc, Ld, LD, lD, ld, le, lE, Le, LE, lf, lF, LF, Lf, LG, Lg, lG, lg, LH, Lh, lh, lH, li, lI, Li, LI, lj, lJ, Lj, LJ, Lk, lK, lk, LK, LL, Ll, ll, lL, lM, LM, Lm, lm, LN, Ln, ln, lN, lo, lO, LO, Lo, lp, lP, Lp, LP, lQ, Lq, LQ, lq, lr, lR, LR, Lr, ls, lS, LS, Ls, Lt, LT, lt, lT, Lu, LU, lU, lu, lV, lv, LV, Lv, lW, lw, LW, Lw, Lx, LX, lX, lx, LY, Ly, ly, lY, lz, lZ, Lz, LZ, m, M, MA, ma, mA, Ma, mb, mB, Mb, MB, mC, MC, Mc, mc, MD, Md, mD, md, Me, ME, me, mE, mF, mf, Mf, MF, mG, mg, MG, Mg, MH, Mh, mH, mh, Mi, MI, mI, mi, mj, mJ, Mj, MJ, mk, mK, MK, Mk, ML, Ml, ml, mL, Mm, MM, mM, mm, MN, Mn, mN, mn, MO, Mo, mO, mo, Mp, mP, mp, MP, mQ, mq, MQ, Mq, Mr, MR, mr, mR, mS, Ms, ms, MS, mt, mT, MT, Mt, MU, Mu, mU, mu, MV, Mv, mV, mv, mw, mW, MW, Mw, mx, mX, MX, Mx, MY, My, mY, my, MZ, Mz, mz, mZ, n, N, NA, Na, na, nA, Nb, nB, nb, NB, nC, nc, Nc, NC, nd, Nd, ND, nD, Ne, NE, ne, nE, NF, Nf, nF, nf, nG, ng, NG, Ng, nh, nH, NH, Nh, NI, Ni, ni, nI, NJ, Nj, nJ, nj, nK, nk, Nk, NK, nl, nL, Nl, NL, NM, Nm, nm, nM, NN, Nn, nN, nn, NO, No, no, nO, Np, NP, nP, np, NQ, Nq, nQ, nq, nr, nR, NR, Nr, NS, Ns, ns, nS, Nt, NT, nT, nt, nU, nu, Nu, NU, NV, Nv, nV, nv, Nw, NW, nW, nw, nX, nx, Nx, NX, ny, nY, NY, Ny, Nz, NZ, nz, nZ, O, o, Oa, OA, oA, oa, Ob, OB, ob, oB, Oc, OC, oc, oC, oD, od, OD, Od, oe, oE, Oe, OE, Of, OF, oF, of, Og, OG, oG, og, oh, oH, OH, Oh, oI, oi, Oi, OI, Oj, OJ, oj, oJ, Ok, OK, ok, oK, ol, oL, OL, Ol, om, oM, Om, OM, On, ON, oN, on, Oo, OO, oO, oo, Op, OP, oP, op, Oq, OQ, oQ, oq, OR, or, oR, Or, oS, os, Os, OS, ot, OT, oT, Ot, Ou, OU, ou, oU, oV, ov, OV, Ov, OW, Ow, ow, oW, OX, Ox, ox, oX, oY, oy, Oy, OY, Oz, oZ, oz, OZ, P, p, pa, pA, Pa, PA, PB, Pb, pb, pB, PC, pC, pc, Pc, Pd, PD, pD, pd, pe, pE, Pe, PE, pf, pF, PF, Pf, PG, pG, Pg, pg, PH, ph, pH, Ph, pI, Pi, PI, Pj, pj, pJ, PJ, PK, pK, pk, Pk, PL, Pl, pL, pl, pm, pM, Pm, PM, Pn, pn, pN, PN, PO, pO, po, Po, PP, Pp, pp, pP, PQ, Pq, pq, pQ, PR, pr, pR, Pr, Ps, pS, ps, PS, pt, Pt, pT, PT, pU, PU, Pu, pu, Pv, PV, pv, pV, pW, pw, PW, Pw, px, PX, pX, Px, Py, PY, py, pY, pZ, Pz, PZ, pz, q, Q, qa, qA, Qa, QA, qB, Qb, QB, qb, qc, Qc, QC, qC, qd, qD, QD, Qd, Qe, QE, qe, qE, Qf, qF, QF, qf, qG, qg, Qg, QG, qh, QH, Qh, qH, Qi, qI, qi, QI, qJ, Qj, QJ, qj, Qk, QK, qk, qK, ql, qL, QL, Ql, qm, Qm, QM, qM, qN, Qn, QN, qn, Qo, QO, qO, qo, qp, qP, QP, Qp, qq, Qq, QQ, qQ, qr, Qr, QR, qR, qs, qS, Qs, QS, QT, Qt, qT, qt, qU, QU, Qu, qu, qv, Qv, QV, qV, Qw, qw, QW, qW, qX, Qx, QX, qx, qy, qY, Qy, QY, QZ, Qz, qz, qZ, r, R, RA, Ra, rA, ra, rB, rb, RB, Rb, Rc, RC, rC, rc, rd, rD, Rd, RD, re, rE, RE, Re, RF, Rf, rf, rF, RG, Rg, rG, rg, rh, rH, Rh, RH, ri, rI, Ri, RI, RJ, rj, rJ, Rj, Rk, RK, rK, rk, Rl, RL, rL, rl, rm, rM, RM, Rm, rn, rN, RN, Rn, RO, Ro, rO, ro, RP, Rp, rp, rP, rq, rQ, RQ, Rq, rr, rR, RR, Rr, rs, rS, Rs, RS, rT, rt, RT, Rt, Ru, RU, ru, rU, Rv, RV, rv, rV, rw, rW, RW, Rw, RX, rX, rx, Rx, Ry, RY, ry, rY, rz, rZ, RZ, Rz, s, S, Sa, SA, sA, sa, Sb, SB, sb, sB, sc, sC, Sc, SC, SD, Sd, sd, sD, se, sE, Se, SE, sF, sf, Sf, SF, SG, Sg, sG, sg, SH, Sh, sH, sh, sI, si, SI, Si, sj, sJ, Sj, SJ, sk, sK, SK, Sk, SL, Sl, sL, sl, Sm, SM, sm, sM, sN, sn, Sn, SN, sO, so, SO, So, SP, Sp, sp, sP, Sq, SQ, sQ, sq, sr, sR, Sr, SR, ss, sS, SS, Ss, st, sT, ST, St, su, sU, SU, Su, SV, Sv, sV, sv, SW, Sw, sW, sw, sX, sx, Sx, SX, Sy, sY, SY, sy, SZ, Sz, sz, sZ, T, t, tA, ta, TA, Ta, TB, Tb, tB, tb, TC, Tc, tc, tC, tD, td, Td, TD, TE, Te, te, tE, tF, tf, TF, Tf, tG, tg, TG, Tg, TH, Th, th, tH, Ti, TI, ti, tI, tJ, tj, Tj, TJ, tK, tk, TK, Tk, Tl, tl, tL, TL, Tm, TM, tm, tM, TN, tN, tn, Tn, tO, to, To, TO, tp, tP, Tp, TP, TQ, Tq, tq, tQ, TR, Tr, tR, tr, tS, ts, Ts, TS, tT, tt, Tt, TT, tU, tu, Tu, TU, tV, tv, Tv, TV, Tw, TW, tW, tw, TX, Tx, tX, tx, ty, tY, Ty, TY, tz, TZ, Tz, tZ, u, U, ua, uA, Ua, UA, uB, Ub, UB, ub, UC, Uc, uc, uC, Ud, UD, ud, uD, ue, uE, Ue, UE, Uf, UF, uF, uf, uG, ug, Ug, UG, uh, uH, UH, Uh, Ui, UI, uI, ui, Uj, UJ, uj, uJ, uk, uK, Uk, UK, ul, uL, UL, Ul, Um, UM, uM, um, Un, UN, un, uN, uo, uO, Uo, UO, up, uP, UP, Up, uQ, Uq, UQ, uq, Ur, UR, ur, uR, Us, US, us, uS, uT, ut, UT, Ut, uu, uU, Uu, UU, uv, uV, Uv, UV, uw, uW, UW, Uw, UX, Ux, ux, uX, UY, Uy, uy, uY, uZ, uz, Uz, UZ, v, V, Va, vA, va, VA, vB, vb, VB, Vb, VC, Vc, vc, vC, Vd, VD, vd, vD, VE, Ve, vE, ve, vF, vf, VF, Vf, VG, Vg, vG, vg, vh, vH, VH, Vh, vi, vI, VI, Vi, VJ, Vj, vj, vJ, Vk, VK, vK, vk, vl, vL, Vl, VL, vm, vM, VM, Vm, VN, Vn, vn, vN, VO, Vo, vO, vo, Vp, vp, vP, VP, vq, vQ, VQ, Vq, VR, Vr, vr, vR, Vs, VS, vS, vs, Vt, VT, vT, vt, vu, vU, VU, Vu, vv, vV, VV, Vv, vW, vw, VW, Vw, vx, vX, Vx, VX, Vy, VY, vy, vY, VZ, Vz, vZ, vz, W, w, Wa, WA, wA, wa, WB, wb, wB, Wb, wc, wC, WC, Wc, WD, Wd, wD, wd, We, WE, wE, we, WF, Wf, wF, wf, wG, wg, WG, Wg, WH, Wh, wh, wH, wI, wi, Wi, WI, wj, wJ, Wj, WJ, WK, Wk, wk, wK, WL, Wl, wl, wL, wm, wM, Wm, WM, wN, wn, Wn, WN, Wo, WO, wO, wo, WP, Wp, wp, wP, WQ, wQ, wq, Wq, wr, wR, Wr, WR, Ws, WS, ws, wS, WT, Wt, wT, wt, Wu, WU, wu, wU, wV, wv, Wv, WV, wW, ww, Ww, WW, wX, wx, Wx, WX, wY, wy, Wy, WY, WZ, Wz, wz, wZ, X, x, Xa, XA, xA, xa, XB, Xb, xB, xb, xC, xc, Xc, XC, xd, xD, Xd, XD, xe, Xe, XE, xE, XF, Xf, xF, xf, XG, Xg, xG, xg, xh, xH, XH, Xh, Xi, XI, xi, xI, xJ, xj, XJ, Xj, xK, xk, Xk, XK, XL, Xl, xL, xl, Xm, XM, xm, xM, xN, xn, XN, Xn, xO, xo, Xo, XO, XP, Xp, xp, xP, XQ, Xq, xq, xQ, xr, xR, Xr, XR, xS, xs, Xs, XS, Xt, xt, XT, xT, Xu, XU, xu, xU, XV, xv, xV, Xv, xW, xw, XW, Xw, xx, xX, XX, Xx, xY, xy, Xy, XY, xZ, xz, Xz, XZ, y, Y, ya, yA, Ya, YA, YB, Yb, yb, yB, YC, Yc, yC, yc, yd, yD, Yd, YD, yE, ye, Ye, YE, yF, YF, Yf, yf, YG, Yg, yG, yg, yh, Yh, YH, yH, yI, yi, Yi, YI, YJ, yj, Yj, yJ, yk, yK, Yk, YK, yL, YL, Yl, yl, Ym, YM, yM, ym, YN, Yn, yn, yN, yO, yo, YO, Yo, yp, yP, YP, Yp, Yq, YQ, yQ, yq, Yr, YR, yR, yr, ys, yS, Ys, YS, yT, yt, YT, Yt, Yu, YU, yU, yu, Yv, YV, yV, yv, yw, yW, Yw, YW, yx, yX, YX, Yx, yY, Yy, yy, YY, yz, yZ, YZ, Yz, z, Z, zA, ZA, za, Za, zB, zb, ZB, Zb, Zc, ZC, zc, zC, Zd, ZD, zd, zD, ze, zE, ZE, Ze, zF, zf, Zf, ZF, Zg, ZG, zg, zG, ZH, Zh, zh, zH, zI, zi, Zi, ZI, zJ, zj, Zj, ZJ, Zk, zK, zk, ZK, zL, zl, Zl, ZL, zm, ZM, Zm, zM, ZN, Zn, zn, zN, Zo, ZO, zO, zo, zp, zP, ZP, Zp, zQ, zq, ZQ, Zq, ZR, Zr, zR, zr, Zs, ZS, zS, zs, zt, zT, Zt, ZT, zu, zU, Zu, ZU, ZV, Zv, zv, zV, ZW, Zw, zW, zw, zx, zX, Zx, ZX, zy, zY, Zy, ZY, ZZ, zz, zZ, Zz
            </forbidwords>
            <allowwords>i, pi, e, ket, depends</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options>allowEmpty, allowEmpty</options>
        </input>
        <input>
            <name>ans2</name>
            <type>algebraic</type>
            <tans>depends</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>
                A, a, aA, Aa, AA, aa, ab, aB, Ab, AB, Ac, AC, aC, ac, aD, AD, Ad, ad, ae, AE, Ae, aE, Af, AF, af, aF, Ag, aG, AG, ag, aH, ah, Ah, AH, ai, aI, AI, Ai, AJ, Aj, aJ, aj, AK, Ak, aK, ak, al, aL, Al, AL, am, aM, Am, AM, AN, An, an, aN, AO, Ao, aO, ao, ap, aP, AP, Ap, aq, aQ, AQ, Aq, ar, aR, AR, Ar, aS, as, AS, As, At, AT, at, aT, Au, AU, au, aU, av, aV, Av, AV, aW, aw, Aw, AW, AX, ax, aX, Ax, AY, Ay, ay, aY, AZ, Az, aZ, az, b, B, ba, bA, BA, Ba, BB, Bb, bB, bb, bc, bC, BC, Bc, BD, Bd, bD, bd, Be, BE, bE, be, bF, bf, Bf, BF, BG, bg, bG, Bg, BH, Bh, bH, bh, bi, bI, BI, Bi, bj, bJ, Bj, BJ, Bk, BK, bk, bK, BL, Bl, bl, bL, bM, bm, BM, Bm, bN, bn, Bn, BN, Bo, BO, bO, bo, BP, Bp, bP, bp, bq, bQ, BQ, Bq, bR, br, Br, BR, bs, bS, Bs, BS, bt, bT, BT, Bt, Bu, BU, bU, bu, Bv, BV, bV, bv, bW, bw, BW, Bw, bX, bx, BX, Bx, by, bY, By, BY, Bz, BZ, bz, bZ, C, c, cA, ca, Ca, CA, cb, cB, CB, Cb, CC, Cc, cc, cC, cd, cD, CD, Cd, CE, Ce, ce, cE, CF, Cf, cF, cf, cg, cG, CG, Cg, Ch, cH, ch, CH, CI, Ci, cI, ci, cj, cJ, CJ, Cj, Ck, cK, ck, CK, Cl, CL, cL, cl, CM, cm, cM, Cm, cn, cN, Cn, CN, cO, co, Co, CO, CP, Cp, cp, cP, CQ, Cq, cq, cQ, cR, cr, CR, Cr, cS, cs, Cs, CS, ct, cT, CT, Ct, cu, cU, Cu, CU, CV, Cv, cv, cV, CW, Cw, cW, cw, cx, cX, Cx, CX, cy, cY, CY, Cy, Cz, cz, cZ, CZ, D, d, DA, Da, da, dA, dB, db, DB, Db, dC, dc, Dc, DC, Dd, DD, dD, dd, dE, de, De, DE, Df, DF, dF, df, DG, Dg, dG, dg, dH, dh, DH, Dh, DI, Di, di, dI, Dj, DJ, dj, dJ, dK, dk, Dk, DK, DL, Dl, dL, dl, Dm, DM, dm, dM, dn, dN, Dn, DN, do, dO, Do, DO, dP, DP, Dp, dp, Dq, DQ, dq, dQ, Dr, DR, dR, dr, ds, dS, Ds, DS, dt, dT, DT, Dt, dU, du, Du, DU, dV, dv, DV, Dv, DW, Dw, dW, dw, Dx, DX, dX, dx, dy, dY, Dy, DY, dz, dZ, DZ, Dz, E, Ea, EA, ea, eA, Eb, EB, eb, eB, ec, eC, EC, Ec, ed, eD, ED, Ed, Ee, EE, ee, eE, ef, eF, EF, Ef, EG, Eg, eg, eG, EH, Eh, eh, eH, ei, eI, EI, Ei, EJ, Ej, ej, eJ, ek, EK, Ek, eK, el, eL, El, EL, Em, EM, em, eM, EN, En, en, eN, Eo, eO, eo, EO, ep, eP, EP, Ep, EQ, Eq, eq, eQ, ER, Er, eR, er, ES, Es, eS, es, et, eT, Et, ET, eu, eU, Eu, EU, eV, ev, EV, Ev, ew, eW, EW, Ew, EX, Ex, ex, eX, EY, Ey, eY, ey, ez, eZ, EZ, Ez, f, F, FA, Fa, fa, fA, FB, Fb, fb, fB, Fc, FC, fC, fc, fD, fd, Fd, FD, fe, fE, FE, Fe, FF, Ff, fF, ff, fg, fG, Fg, FG, Fh, FH, fH, fh, FI, Fi, fI, fi, fJ, fj, FJ, Fj, Fk, fk, FK, fK, fL, fl, FL, Fl, fM, fm, FM, Fm, Fn, FN, fN, fn, Fo, FO, fO, fo, FP, Fp, fP, fp, fq, fQ, Fq, FQ, fr, Fr, FR, fR, Fs, FS, fs, fS, FT, Ft, ft, fT, fU, fu, Fu, FU, fV, fv, Fv, FV, fW, fw, FW, Fw, fX, fx, Fx, FX, FY, Fy, fy, fY, Fz, FZ, fz, fZ, G, g, gA, ga, Ga, GA, Gb, GB, gb, gB, Gc, GC, gc, gC, GD, Gd, gD, gd, ge, gE, Ge, GE, gF, gf, Gf, GF, Gg, GG, gG, gg, gH, gh, Gh, GH, Gi, GI, gI, gi, Gj, GJ, gJ, gj, gK, gk, Gk, GK, Gl, gl, GL, gL, gm, gM, GM, Gm, gn, gN, Gn, GN, GO, Go, gO, go, GP, Gp, gp, gP, GQ, gq, gQ, Gq, gR, gr, GR, Gr, Gs, gS, gs, GS, Gt, GT, gt, gT, GU, gu, Gu, gU, gv, gV, Gv, GV, gW, gw, GW, Gw, gx, gX, GX, Gx, gY, gy, Gy, GY, GZ, Gz, gZ, gz, H, h, ha, hA, Ha, HA, hb, hB, HB, Hb, Hc, HC, hC, hc, Hd, HD, hD, hd, HE, he, hE, He, hf, hF, HF, Hf, hG, Hg, hg, HG, Hh, HH, hh, hH, hI, Hi, hi, HI, HJ, Hj, hj, hJ, Hk, HK, hk, hK, hl, hL, HL, Hl, Hm, HM, hM, hm, hN, hn, HN, Hn, ho, hO, HO, Ho, Hp, HP, hP, hp, Hq, HQ, hQ, hq, hr, hR, Hr, HR, hs, hS, Hs, HS, HT, Ht, hT, ht, Hu, HU, hU, hu, hv, Hv, hV, HV, hW, hw, HW, Hw, hX, hx, HX, Hx, hy, hY, Hy, HY, hZ, HZ, Hz, hz, I, Ia, IA, ia, iA, ib, iB, Ib, IB, iC, ic, IC, Ic, ID, Id, id, iD, Ie, IE, ie, iE, if, If, iF, IF, iG, ig, IG, Ig, IH, Ih, iH, ih, II, Ii, iI, ii, ij, Ij, IJ, iJ, Ik, IK, iK, ik, IL, Il, il, iL, iM, im, IM, Im, IN, in, In, iN, iO, io, IO, Io, ip, iP, IP, Ip, IQ, Iq, iQ, iq, IR, Ir, ir, iR, iS, is, IS, Is, it, iT, It, IT, Iu, IU, iu, iU, Iv, IV, iv, iV, iW, iw, Iw, IW, ix, iX, Ix, IX, iy, iY, IY, Iy, iZ, iz, Iz, IZ, j, J, JA, Ja, ja, jA, JB, Jb, jB, jb, jC, jc, Jc, JC, jd, jD, Jd, JD, JE, Je, je, jE, JF, Jf, jF, jf, jG, jg, JG, Jg, jh, jH, Jh, JH, JI, ji, jI, Ji, Jj, JJ, jJ, jj, jk, jK, Jk, JK, JL, Jl, jL, jl, Jm, JM, jM, jm, jN, jn, Jn, JN, jo, jO, Jo, JO, jP, jp, Jp, JP, jq, jQ, JQ, Jq, Jr, JR, jr, jR, Js, JS, js, jS, jt, jT, JT, Jt, jU, ju, Ju, JU, Jv, JV, jV, jv, JW, Jw, jW, jw, jX, jx, JX, Jx, jy, jY, JY, Jy, jz, jZ, Jz, JZ, K, k, kA, ka, Ka, KA, Kb, KB, kb, kB, Kc, KC, kc, kC, kd, kD, KD, Kd, ke, kE, Ke, KE, Kf, KF, kF, kf, Kg, KG, kG, kg, kH, kh, KH, Kh, kI, ki, KI, Ki, kj, kJ, Kj, KJ, Kk, KK, kk, kK, kl, kL, Kl, KL, Km, KM, km, kM, KN, kn, kN, Kn, ko, kO, KO, Ko, kP, kp, KP, Kp, kQ, KQ, kq, Kq, kR, kr, Kr, KR, Ks, KS, kS, ks, Kt, KT, kt, kT, ku, kU, Ku, KU, kv, kV, KV, Kv, KW, Kw, kW, kw, Kx, KX, kx, kX, ky, kY, KY, Ky, kz, kZ, KZ, Kz, L, l, la, lA, LA, La, lb, lB, LB, Lb, Lc, LC, lC, lc, Ld, LD, lD, ld, le, lE, Le, LE, lf, lF, LF, Lf, LG, Lg, lG, lg, LH, Lh, lh, lH, li, lI, Li, LI, lj, lJ, Lj, LJ, Lk, lK, lk, LK, LL, Ll, ll, lL, lM, LM, Lm, lm, LN, Ln, ln, lN, lo, lO, LO, Lo, lp, lP, Lp, LP, lQ, Lq, LQ, lq, lr, lR, LR, Lr, ls, lS, LS, Ls, Lt, LT, lt, lT, Lu, LU, lU, lu, lV, lv, LV, Lv, lW, lw, LW, Lw, Lx, LX, lX, lx, LY, Ly, ly, lY, lz, lZ, Lz, LZ, m, M, MA, ma, mA, Ma, mb, mB, Mb, MB, mC, MC, Mc, mc, MD, Md, mD, md, Me, ME, me, mE, mF, mf, Mf, MF, mG, mg, MG, Mg, MH, Mh, mH, mh, Mi, MI, mI, mi, mj, mJ, Mj, MJ, mk, mK, MK, Mk, ML, Ml, ml, mL, Mm, MM, mM, mm, MN, Mn, mN, mn, MO, Mo, mO, mo, Mp, mP, mp, MP, mQ, mq, MQ, Mq, Mr, MR, mr, mR, mS, Ms, ms, MS, mt, mT, MT, Mt, MU, Mu, mU, mu, MV, Mv, mV, mv, mw, mW, MW, Mw, mx, mX, MX, Mx, MY, My, mY, my, MZ, Mz, mz, mZ, n, N, NA, Na, na, nA, Nb, nB, nb, NB, nC, nc, Nc, NC, nd, Nd, ND, nD, Ne, NE, ne, nE, NF, Nf, nF, nf, nG, ng, NG, Ng, nh, nH, NH, Nh, NI, Ni, ni, nI, NJ, Nj, nJ, nj, nK, nk, Nk, NK, nl, nL, Nl, NL, NM, Nm, nm, nM, NN, Nn, nN, nn, NO, No, no, nO, Np, NP, nP, np, NQ, Nq, nQ, nq, nr, nR, NR, Nr, NS, Ns, ns, nS, Nt, NT, nT, nt, nU, nu, Nu, NU, NV, Nv, nV, nv, Nw, NW, nW, nw, nX, nx, Nx, NX, ny, nY, NY, Ny, Nz, NZ, nz, nZ, O, o, Oa, OA, oA, oa, Ob, OB, ob, oB, Oc, OC, oc, oC, oD, od, OD, Od, oe, oE, Oe, OE, Of, OF, oF, of, Og, OG, oG, og, oh, oH, OH, Oh, oI, oi, Oi, OI, Oj, OJ, oj, oJ, Ok, OK, ok, oK, ol, oL, OL, Ol, om, oM, Om, OM, On, ON, oN, on, Oo, OO, oO, oo, Op, OP, oP, op, Oq, OQ, oQ, oq, OR, or, oR, Or, oS, os, Os, OS, ot, OT, oT, Ot, Ou, OU, ou, oU, oV, ov, OV, Ov, OW, Ow, ow, oW, OX, Ox, ox, oX, oY, oy, Oy, OY, Oz, oZ, oz, OZ, P, p, pa, pA, Pa, PA, PB, Pb, pb, pB, PC, pC, pc, Pc, Pd, PD, pD, pd, pe, pE, Pe, PE, pf, pF, PF, Pf, PG, pG, Pg, pg, PH, ph, pH, Ph, pI, Pi, PI, Pj, pj, pJ, PJ, PK, pK, pk, Pk, PL, Pl, pL, pl, pm, pM, Pm, PM, Pn, pn, pN, PN, PO, pO, po, Po, PP, Pp, pp, pP, PQ, Pq, pq, pQ, PR, pr, pR, Pr, Ps, pS, ps, PS, pt, Pt, pT, PT, pU, PU, Pu, pu, Pv, PV, pv, pV, pW, pw, PW, Pw, px, PX, pX, Px, Py, PY, py, pY, pZ, Pz, PZ, pz, q, Q, qa, qA, Qa, QA, qB, Qb, QB, qb, qc, Qc, QC, qC, qd, qD, QD, Qd, Qe, QE, qe, qE, Qf, qF, QF, qf, qG, qg, Qg, QG, qh, QH, Qh, qH, Qi, qI, qi, QI, qJ, Qj, QJ, qj, Qk, QK, qk, qK, ql, qL, QL, Ql, qm, Qm, QM, qM, qN, Qn, QN, qn, Qo, QO, qO, qo, qp, qP, QP, Qp, qq, Qq, QQ, qQ, qr, Qr, QR, qR, qs, qS, Qs, QS, QT, Qt, qT, qt, qU, QU, Qu, qu, qv, Qv, QV, qV, Qw, qw, QW, qW, qX, Qx, QX, qx, qy, qY, Qy, QY, QZ, Qz, qz, qZ, r, R, RA, Ra, rA, ra, rB, rb, RB, Rb, Rc, RC, rC, rc, rd, rD, Rd, RD, re, rE, RE, Re, RF, Rf, rf, rF, RG, Rg, rG, rg, rh, rH, Rh, RH, ri, rI, Ri, RI, RJ, rj, rJ, Rj, Rk, RK, rK, rk, Rl, RL, rL, rl, rm, rM, RM, Rm, rn, rN, RN, Rn, RO, Ro, rO, ro, RP, Rp, rp, rP, rq, rQ, RQ, Rq, rr, rR, RR, Rr, rs, rS, Rs, RS, rT, rt, RT, Rt, Ru, RU, ru, rU, Rv, RV, rv, rV, rw, rW, RW, Rw, RX, rX, rx, Rx, Ry, RY, ry, rY, rz, rZ, RZ, Rz, s, S, Sa, SA, sA, sa, Sb, SB, sb, sB, sc, sC, Sc, SC, SD, Sd, sd, sD, se, sE, Se, SE, sF, sf, Sf, SF, SG, Sg, sG, sg, SH, Sh, sH, sh, sI, si, SI, Si, sj, sJ, Sj, SJ, sk, sK, SK, Sk, SL, Sl, sL, sl, Sm, SM, sm, sM, sN, sn, Sn, SN, sO, so, SO, So, SP, Sp, sp, sP, Sq, SQ, sQ, sq, sr, sR, Sr, SR, ss, sS, SS, Ss, st, sT, ST, St, su, sU, SU, Su, SV, Sv, sV, sv, SW, Sw, sW, sw, sX, sx, Sx, SX, Sy, sY, SY, sy, SZ, Sz, sz, sZ, T, t, tA, ta, TA, Ta, TB, Tb, tB, tb, TC, Tc, tc, tC, tD, td, Td, TD, TE, Te, te, tE, tF, tf, TF, Tf, tG, tg, TG, Tg, TH, Th, th, tH, Ti, TI, ti, tI, tJ, tj, Tj, TJ, tK, tk, TK, Tk, Tl, tl, tL, TL, Tm, TM, tm, tM, TN, tN, tn, Tn, tO, to, To, TO, tp, tP, Tp, TP, TQ, Tq, tq, tQ, TR, Tr, tR, tr, tS, ts, Ts, TS, tT, tt, Tt, TT, tU, tu, Tu, TU, tV, tv, Tv, TV, Tw, TW, tW, tw, TX, Tx, tX, tx, ty, tY, Ty, TY, tz, TZ, Tz, tZ, u, U, ua, uA, Ua, UA, uB, Ub, UB, ub, UC, Uc, uc, uC, Ud, UD, ud, uD, ue, uE, Ue, UE, Uf, UF, uF, uf, uG, ug, Ug, UG, uh, uH, UH, Uh, Ui, UI, uI, ui, Uj, UJ, uj, uJ, uk, uK, Uk, UK, ul, uL, UL, Ul, Um, UM, uM, um, Un, UN, un, uN, uo, uO, Uo, UO, up, uP, UP, Up, uQ, Uq, UQ, uq, Ur, UR, ur, uR, Us, US, us, uS, uT, ut, UT, Ut, uu, uU, Uu, UU, uv, uV, Uv, UV, uw, uW, UW, Uw, UX, Ux, ux, uX, UY, Uy, uy, uY, uZ, uz, Uz, UZ, v, V, Va, vA, va, VA, vB, vb, VB, Vb, VC, Vc, vc, vC, Vd, VD, vd, vD, VE, Ve, vE, ve, vF, vf, VF, Vf, VG, Vg, vG, vg, vh, vH, VH, Vh, vi, vI, VI, Vi, VJ, Vj, vj, vJ, Vk, VK, vK, vk, vl, vL, Vl, VL, vm, vM, VM, Vm, VN, Vn, vn, vN, VO, Vo, vO, vo, Vp, vp, vP, VP, vq, vQ, VQ, Vq, VR, Vr, vr, vR, Vs, VS, vS, vs, Vt, VT, vT, vt, vu, vU, VU, Vu, vv, vV, VV, Vv, vW, vw, VW, Vw, vx, vX, Vx, VX, Vy, VY, vy, vY, VZ, Vz, vZ, vz, W, w, Wa, WA, wA, wa, WB, wb, wB, Wb, wc, wC, WC, Wc, WD, Wd, wD, wd, We, WE, wE, we, WF, Wf, wF, wf, wG, wg, WG, Wg, WH, Wh, wh, wH, wI, wi, Wi, WI, wj, wJ, Wj, WJ, WK, Wk, wk, wK, WL, Wl, wl, wL, wm, wM, Wm, WM, wN, wn, Wn, WN, Wo, WO, wO, wo, WP, Wp, wp, wP, WQ, wQ, wq, Wq, wr, wR, Wr, WR, Ws, WS, ws, wS, WT, Wt, wT, wt, Wu, WU, wu, wU, wV, wv, Wv, WV, wW, ww, Ww, WW, wX, wx, Wx, WX, wY, wy, Wy, WY, WZ, Wz, wz, wZ, X, x, Xa, XA, xA, xa, XB, Xb, xB, xb, xC, xc, Xc, XC, xd, xD, Xd, XD, xe, Xe, XE, xE, XF, Xf, xF, xf, XG, Xg, xG, xg, xh, xH, XH, Xh, Xi, XI, xi, xI, xJ, xj, XJ, Xj, xK, xk, Xk, XK, XL, Xl, xL, xl, Xm, XM, xm, xM, xN, xn, XN, Xn, xO, xo, Xo, XO, XP, Xp, xp, xP, XQ, Xq, xq, xQ, xr, xR, Xr, XR, xS, xs, Xs, XS, Xt, xt, XT, xT, Xu, XU, xu, xU, XV, xv, xV, Xv, xW, xw, XW, Xw, xx, xX, XX, Xx, xY, xy, Xy, XY, xZ, xz, Xz, XZ, y, Y, ya, yA, Ya, YA, YB, Yb, yb, yB, YC, Yc, yC, yc, yd, yD, Yd, YD, yE, ye, Ye, YE, yF, YF, Yf, yf, YG, Yg, yG, yg, yh, Yh, YH, yH, yI, yi, Yi, YI, YJ, yj, Yj, yJ, yk, yK, Yk, YK, yL, YL, Yl, yl, Ym, YM, yM, ym, YN, Yn, yn, yN, yO, yo, YO, Yo, yp, yP, YP, Yp, Yq, YQ, yQ, yq, Yr, YR, yR, yr, ys, yS, Ys, YS, yT, yt, YT, Yt, Yu, YU, yU, yu, Yv, YV, yV, yv, yw, yW, Yw, YW, yx, yX, YX, Yx, yY, Yy, yy, YY, yz, yZ, YZ, Yz, z, Z, zA, ZA, za, Za, zB, zb, ZB, Zb, Zc, ZC, zc, zC, Zd, ZD, zd, zD, ze, zE, ZE, Ze, zF, zf, Zf, ZF, Zg, ZG, zg, zG, ZH, Zh, zh, zH, zI, zi, Zi, ZI, zJ, zj, Zj, ZJ, Zk, zK, zk, ZK, zL, zl, Zl, ZL, zm, ZM, Zm, zM, ZN, Zn, zn, zN, Zo, ZO, zO, zo, zp, zP, ZP, Zp, zQ, zq, ZQ, Zq, ZR, Zr, zR, zr, Zs, ZS, zS, zs, zt, zT, Zt, ZT, zu, zU, Zu, ZU, ZV, Zv, zv, zV, ZW, Zw, zW, zw, zx, zX, Zx, ZX, zy, zY, Zy, ZY, ZZ, zz, zZ, Zz
            </forbidwords>
            <allowwords>i, pi, e, ket, depends</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options>allowEmpty, allowEmpty</options>
        </input>
        <input>
            <name>ans3</name>
            <type>algebraic</type>
            <tans>(ket(00)+ket(01)-ket(10)-ket(11))/2</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>
                A, a, aA, Aa, AA, aa, ab, aB, Ab, AB, Ac, AC, aC, ac, aD, AD, Ad, ad, ae, AE, Ae, aE, Af, AF, af, aF, Ag, aG, AG, ag, aH, ah, Ah, AH, ai, aI, AI, Ai, AJ, Aj, aJ, aj, AK, Ak, aK, ak, al, aL, Al, AL, am, aM, Am, AM, AN, An, an, aN, AO, Ao, aO, ao, ap, aP, AP, Ap, aq, aQ, AQ, Aq, ar, aR, AR, Ar, aS, as, AS, As, At, AT, at, aT, Au, AU, au, aU, av, aV, Av, AV, aW, aw, Aw, AW, AX, ax, aX, Ax, AY, Ay, ay, aY, AZ, Az, aZ, az, b, B, ba, bA, BA, Ba, BB, Bb, bB, bb, bc, bC, BC, Bc, BD, Bd, bD, bd, Be, BE, bE, be, bF, bf, Bf, BF, BG, bg, bG, Bg, BH, Bh, bH, bh, bi, bI, BI, Bi, bj, bJ, Bj, BJ, Bk, BK, bk, bK, BL, Bl, bl, bL, bM, bm, BM, Bm, bN, bn, Bn, BN, Bo, BO, bO, bo, BP, Bp, bP, bp, bq, bQ, BQ, Bq, bR, br, Br, BR, bs, bS, Bs, BS, bt, bT, BT, Bt, Bu, BU, bU, bu, Bv, BV, bV, bv, bW, bw, BW, Bw, bX, bx, BX, Bx, by, bY, By, BY, Bz, BZ, bz, bZ, C, c, cA, ca, Ca, CA, cb, cB, CB, Cb, CC, Cc, cc, cC, cd, cD, CD, Cd, CE, Ce, ce, cE, CF, Cf, cF, cf, cg, cG, CG, Cg, Ch, cH, ch, CH, CI, Ci, cI, ci, cj, cJ, CJ, Cj, Ck, cK, ck, CK, Cl, CL, cL, cl, CM, cm, cM, Cm, cn, cN, Cn, CN, cO, co, Co, CO, CP, Cp, cp, cP, CQ, Cq, cq, cQ, cR, cr, CR, Cr, cS, cs, Cs, CS, ct, cT, CT, Ct, cu, cU, Cu, CU, CV, Cv, cv, cV, CW, Cw, cW, cw, cx, cX, Cx, CX, cy, cY, CY, Cy, Cz, cz, cZ, CZ, D, d, DA, Da, da, dA, dB, db, DB, Db, dC, dc, Dc, DC, Dd, DD, dD, dd, dE, de, De, DE, Df, DF, dF, df, DG, Dg, dG, dg, dH, dh, DH, Dh, DI, Di, di, dI, Dj, DJ, dj, dJ, dK, dk, Dk, DK, DL, Dl, dL, dl, Dm, DM, dm, dM, dn, dN, Dn, DN, do, dO, Do, DO, dP, DP, Dp, dp, Dq, DQ, dq, dQ, Dr, DR, dR, dr, ds, dS, Ds, DS, dt, dT, DT, Dt, dU, du, Du, DU, dV, dv, DV, Dv, DW, Dw, dW, dw, Dx, DX, dX, dx, dy, dY, Dy, DY, dz, dZ, DZ, Dz, E, Ea, EA, ea, eA, Eb, EB, eb, eB, ec, eC, EC, Ec, ed, eD, ED, Ed, Ee, EE, ee, eE, ef, eF, EF, Ef, EG, Eg, eg, eG, EH, Eh, eh, eH, ei, eI, EI, Ei, EJ, Ej, ej, eJ, ek, EK, Ek, eK, el, eL, El, EL, Em, EM, em, eM, EN, En, en, eN, Eo, eO, eo, EO, ep, eP, EP, Ep, EQ, Eq, eq, eQ, ER, Er, eR, er, ES, Es, eS, es, et, eT, Et, ET, eu, eU, Eu, EU, eV, ev, EV, Ev, ew, eW, EW, Ew, EX, Ex, ex, eX, EY, Ey, eY, ey, ez, eZ, EZ, Ez, f, F, FA, Fa, fa, fA, FB, Fb, fb, fB, Fc, FC, fC, fc, fD, fd, Fd, FD, fe, fE, FE, Fe, FF, Ff, fF, ff, fg, fG, Fg, FG, Fh, FH, fH, fh, FI, Fi, fI, fi, fJ, fj, FJ, Fj, Fk, fk, FK, fK, fL, fl, FL, Fl, fM, fm, FM, Fm, Fn, FN, fN, fn, Fo, FO, fO, fo, FP, Fp, fP, fp, fq, fQ, Fq, FQ, fr, Fr, FR, fR, Fs, FS, fs, fS, FT, Ft, ft, fT, fU, fu, Fu, FU, fV, fv, Fv, FV, fW, fw, FW, Fw, fX, fx, Fx, FX, FY, Fy, fy, fY, Fz, FZ, fz, fZ, G, g, gA, ga, Ga, GA, Gb, GB, gb, gB, Gc, GC, gc, gC, GD, Gd, gD, gd, ge, gE, Ge, GE, gF, gf, Gf, GF, Gg, GG, gG, gg, gH, gh, Gh, GH, Gi, GI, gI, gi, Gj, GJ, gJ, gj, gK, gk, Gk, GK, Gl, gl, GL, gL, gm, gM, GM, Gm, gn, gN, Gn, GN, GO, Go, gO, go, GP, Gp, gp, gP, GQ, gq, gQ, Gq, gR, gr, GR, Gr, Gs, gS, gs, GS, Gt, GT, gt, gT, GU, gu, Gu, gU, gv, gV, Gv, GV, gW, gw, GW, Gw, gx, gX, GX, Gx, gY, gy, Gy, GY, GZ, Gz, gZ, gz, H, h, ha, hA, Ha, HA, hb, hB, HB, Hb, Hc, HC, hC, hc, Hd, HD, hD, hd, HE, he, hE, He, hf, hF, HF, Hf, hG, Hg, hg, HG, Hh, HH, hh, hH, hI, Hi, hi, HI, HJ, Hj, hj, hJ, Hk, HK, hk, hK, hl, hL, HL, Hl, Hm, HM, hM, hm, hN, hn, HN, Hn, ho, hO, HO, Ho, Hp, HP, hP, hp, Hq, HQ, hQ, hq, hr, hR, Hr, HR, hs, hS, Hs, HS, HT, Ht, hT, ht, Hu, HU, hU, hu, hv, Hv, hV, HV, hW, hw, HW, Hw, hX, hx, HX, Hx, hy, hY, Hy, HY, hZ, HZ, Hz, hz, I, Ia, IA, ia, iA, ib, iB, Ib, IB, iC, ic, IC, Ic, ID, Id, id, iD, Ie, IE, ie, iE, if, If, iF, IF, iG, ig, IG, Ig, IH, Ih, iH, ih, II, Ii, iI, ii, ij, Ij, IJ, iJ, Ik, IK, iK, ik, IL, Il, il, iL, iM, im, IM, Im, IN, in, In, iN, iO, io, IO, Io, ip, iP, IP, Ip, IQ, Iq, iQ, iq, IR, Ir, ir, iR, iS, is, IS, Is, it, iT, It, IT, Iu, IU, iu, iU, Iv, IV, iv, iV, iW, iw, Iw, IW, ix, iX, Ix, IX, iy, iY, IY, Iy, iZ, iz, Iz, IZ, j, J, JA, Ja, ja, jA, JB, Jb, jB, jb, jC, jc, Jc, JC, jd, jD, Jd, JD, JE, Je, je, jE, JF, Jf, jF, jf, jG, jg, JG, Jg, jh, jH, Jh, JH, JI, ji, jI, Ji, Jj, JJ, jJ, jj, jk, jK, Jk, JK, JL, Jl, jL, jl, Jm, JM, jM, jm, jN, jn, Jn, JN, jo, jO, Jo, JO, jP, jp, Jp, JP, jq, jQ, JQ, Jq, Jr, JR, jr, jR, Js, JS, js, jS, jt, jT, JT, Jt, jU, ju, Ju, JU, Jv, JV, jV, jv, JW, Jw, jW, jw, jX, jx, JX, Jx, jy, jY, JY, Jy, jz, jZ, Jz, JZ, K, k, kA, ka, Ka, KA, Kb, KB, kb, kB, Kc, KC, kc, kC, kd, kD, KD, Kd, ke, kE, Ke, KE, Kf, KF, kF, kf, Kg, KG, kG, kg, kH, kh, KH, Kh, kI, ki, KI, Ki, kj, kJ, Kj, KJ, Kk, KK, kk, kK, kl, kL, Kl, KL, Km, KM, km, kM, KN, kn, kN, Kn, ko, kO, KO, Ko, kP, kp, KP, Kp, kQ, KQ, kq, Kq, kR, kr, Kr, KR, Ks, KS, kS, ks, Kt, KT, kt, kT, ku, kU, Ku, KU, kv, kV, KV, Kv, KW, Kw, kW, kw, Kx, KX, kx, kX, ky, kY, KY, Ky, kz, kZ, KZ, Kz, L, l, la, lA, LA, La, lb, lB, LB, Lb, Lc, LC, lC, lc, Ld, LD, lD, ld, le, lE, Le, LE, lf, lF, LF, Lf, LG, Lg, lG, lg, LH, Lh, lh, lH, li, lI, Li, LI, lj, lJ, Lj, LJ, Lk, lK, lk, LK, LL, Ll, ll, lL, lM, LM, Lm, lm, LN, Ln, ln, lN, lo, lO, LO, Lo, lp, lP, Lp, LP, lQ, Lq, LQ, lq, lr, lR, LR, Lr, ls, lS, LS, Ls, Lt, LT, lt, lT, Lu, LU, lU, lu, lV, lv, LV, Lv, lW, lw, LW, Lw, Lx, LX, lX, lx, LY, Ly, ly, lY, lz, lZ, Lz, LZ, m, M, MA, ma, mA, Ma, mb, mB, Mb, MB, mC, MC, Mc, mc, MD, Md, mD, md, Me, ME, me, mE, mF, mf, Mf, MF, mG, mg, MG, Mg, MH, Mh, mH, mh, Mi, MI, mI, mi, mj, mJ, Mj, MJ, mk, mK, MK, Mk, ML, Ml, ml, mL, Mm, MM, mM, mm, MN, Mn, mN, mn, MO, Mo, mO, mo, Mp, mP, mp, MP, mQ, mq, MQ, Mq, Mr, MR, mr, mR, mS, Ms, ms, MS, mt, mT, MT, Mt, MU, Mu, mU, mu, MV, Mv, mV, mv, mw, mW, MW, Mw, mx, mX, MX, Mx, MY, My, mY, my, MZ, Mz, mz, mZ, n, N, NA, Na, na, nA, Nb, nB, nb, NB, nC, nc, Nc, NC, nd, Nd, ND, nD, Ne, NE, ne, nE, NF, Nf, nF, nf, nG, ng, NG, Ng, nh, nH, NH, Nh, NI, Ni, ni, nI, NJ, Nj, nJ, nj, nK, nk, Nk, NK, nl, nL, Nl, NL, NM, Nm, nm, nM, NN, Nn, nN, nn, NO, No, no, nO, Np, NP, nP, np, NQ, Nq, nQ, nq, nr, nR, NR, Nr, NS, Ns, ns, nS, Nt, NT, nT, nt, nU, nu, Nu, NU, NV, Nv, nV, nv, Nw, NW, nW, nw, nX, nx, Nx, NX, ny, nY, NY, Ny, Nz, NZ, nz, nZ, O, o, Oa, OA, oA, oa, Ob, OB, ob, oB, Oc, OC, oc, oC, oD, od, OD, Od, oe, oE, Oe, OE, Of, OF, oF, of, Og, OG, oG, og, oh, oH, OH, Oh, oI, oi, Oi, OI, Oj, OJ, oj, oJ, Ok, OK, ok, oK, ol, oL, OL, Ol, om, oM, Om, OM, On, ON, oN, on, Oo, OO, oO, oo, Op, OP, oP, op, Oq, OQ, oQ, oq, OR, or, oR, Or, oS, os, Os, OS, ot, OT, oT, Ot, Ou, OU, ou, oU, oV, ov, OV, Ov, OW, Ow, ow, oW, OX, Ox, ox, oX, oY, oy, Oy, OY, Oz, oZ, oz, OZ, P, p, pa, pA, Pa, PA, PB, Pb, pb, pB, PC, pC, pc, Pc, Pd, PD, pD, pd, pe, pE, Pe, PE, pf, pF, PF, Pf, PG, pG, Pg, pg, PH, ph, pH, Ph, pI, Pi, PI, Pj, pj, pJ, PJ, PK, pK, pk, Pk, PL, Pl, pL, pl, pm, pM, Pm, PM, Pn, pn, pN, PN, PO, pO, po, Po, PP, Pp, pp, pP, PQ, Pq, pq, pQ, PR, pr, pR, Pr, Ps, pS, ps, PS, pt, Pt, pT, PT, pU, PU, Pu, pu, Pv, PV, pv, pV, pW, pw, PW, Pw, px, PX, pX, Px, Py, PY, py, pY, pZ, Pz, PZ, pz, q, Q, qa, qA, Qa, QA, qB, Qb, QB, qb, qc, Qc, QC, qC, qd, qD, QD, Qd, Qe, QE, qe, qE, Qf, qF, QF, qf, qG, qg, Qg, QG, qh, QH, Qh, qH, Qi, qI, qi, QI, qJ, Qj, QJ, qj, Qk, QK, qk, qK, ql, qL, QL, Ql, qm, Qm, QM, qM, qN, Qn, QN, qn, Qo, QO, qO, qo, qp, qP, QP, Qp, qq, Qq, QQ, qQ, qr, Qr, QR, qR, qs, qS, Qs, QS, QT, Qt, qT, qt, qU, QU, Qu, qu, qv, Qv, QV, qV, Qw, qw, QW, qW, qX, Qx, QX, qx, qy, qY, Qy, QY, QZ, Qz, qz, qZ, r, R, RA, Ra, rA, ra, rB, rb, RB, Rb, Rc, RC, rC, rc, rd, rD, Rd, RD, re, rE, RE, Re, RF, Rf, rf, rF, RG, Rg, rG, rg, rh, rH, Rh, RH, ri, rI, Ri, RI, RJ, rj, rJ, Rj, Rk, RK, rK, rk, Rl, RL, rL, rl, rm, rM, RM, Rm, rn, rN, RN, Rn, RO, Ro, rO, ro, RP, Rp, rp, rP, rq, rQ, RQ, Rq, rr, rR, RR, Rr, rs, rS, Rs, RS, rT, rt, RT, Rt, Ru, RU, ru, rU, Rv, RV, rv, rV, rw, rW, RW, Rw, RX, rX, rx, Rx, Ry, RY, ry, rY, rz, rZ, RZ, Rz, s, S, Sa, SA, sA, sa, Sb, SB, sb, sB, sc, sC, Sc, SC, SD, Sd, sd, sD, se, sE, Se, SE, sF, sf, Sf, SF, SG, Sg, sG, sg, SH, Sh, sH, sh, sI, si, SI, Si, sj, sJ, Sj, SJ, sk, sK, SK, Sk, SL, Sl, sL, sl, Sm, SM, sm, sM, sN, sn, Sn, SN, sO, so, SO, So, SP, Sp, sp, sP, Sq, SQ, sQ, sq, sr, sR, Sr, SR, ss, sS, SS, Ss, st, sT, ST, St, su, sU, SU, Su, SV, Sv, sV, sv, SW, Sw, sW, sw, sX, sx, Sx, SX, Sy, sY, SY, sy, SZ, Sz, sz, sZ, T, t, tA, ta, TA, Ta, TB, Tb, tB, tb, TC, Tc, tc, tC, tD, td, Td, TD, TE, Te, te, tE, tF, tf, TF, Tf, tG, tg, TG, Tg, TH, Th, th, tH, Ti, TI, ti, tI, tJ, tj, Tj, TJ, tK, tk, TK, Tk, Tl, tl, tL, TL, Tm, TM, tm, tM, TN, tN, tn, Tn, tO, to, To, TO, tp, tP, Tp, TP, TQ, Tq, tq, tQ, TR, Tr, tR, tr, tS, ts, Ts, TS, tT, tt, Tt, TT, tU, tu, Tu, TU, tV, tv, Tv, TV, Tw, TW, tW, tw, TX, Tx, tX, tx, ty, tY, Ty, TY, tz, TZ, Tz, tZ, u, U, ua, uA, Ua, UA, uB, Ub, UB, ub, UC, Uc, uc, uC, Ud, UD, ud, uD, ue, uE, Ue, UE, Uf, UF, uF, uf, uG, ug, Ug, UG, uh, uH, UH, Uh, Ui, UI, uI, ui, Uj, UJ, uj, uJ, uk, uK, Uk, UK, ul, uL, UL, Ul, Um, UM, uM, um, Un, UN, un, uN, uo, uO, Uo, UO, up, uP, UP, Up, uQ, Uq, UQ, uq, Ur, UR, ur, uR, Us, US, us, uS, uT, ut, UT, Ut, uu, uU, Uu, UU, uv, uV, Uv, UV, uw, uW, UW, Uw, UX, Ux, ux, uX, UY, Uy, uy, uY, uZ, uz, Uz, UZ, v, V, Va, vA, va, VA, vB, vb, VB, Vb, VC, Vc, vc, vC, Vd, VD, vd, vD, VE, Ve, vE, ve, vF, vf, VF, Vf, VG, Vg, vG, vg, vh, vH, VH, Vh, vi, vI, VI, Vi, VJ, Vj, vj, vJ, Vk, VK, vK, vk, vl, vL, Vl, VL, vm, vM, VM, Vm, VN, Vn, vn, vN, VO, Vo, vO, vo, Vp, vp, vP, VP, vq, vQ, VQ, Vq, VR, Vr, vr, vR, Vs, VS, vS, vs, Vt, VT, vT, vt, vu, vU, VU, Vu, vv, vV, VV, Vv, vW, vw, VW, Vw, vx, vX, Vx, VX, Vy, VY, vy, vY, VZ, Vz, vZ, vz, W, w, Wa, WA, wA, wa, WB, wb, wB, Wb, wc, wC, WC, Wc, WD, Wd, wD, wd, We, WE, wE, we, WF, Wf, wF, wf, wG, wg, WG, Wg, WH, Wh, wh, wH, wI, wi, Wi, WI, wj, wJ, Wj, WJ, WK, Wk, wk, wK, WL, Wl, wl, wL, wm, wM, Wm, WM, wN, wn, Wn, WN, Wo, WO, wO, wo, WP, Wp, wp, wP, WQ, wQ, wq, Wq, wr, wR, Wr, WR, Ws, WS, ws, wS, WT, Wt, wT, wt, Wu, WU, wu, wU, wV, wv, Wv, WV, wW, ww, Ww, WW, wX, wx, Wx, WX, wY, wy, Wy, WY, WZ, Wz, wz, wZ, X, x, Xa, XA, xA, xa, XB, Xb, xB, xb, xC, xc, Xc, XC, xd, xD, Xd, XD, xe, Xe, XE, xE, XF, Xf, xF, xf, XG, Xg, xG, xg, xh, xH, XH, Xh, Xi, XI, xi, xI, xJ, xj, XJ, Xj, xK, xk, Xk, XK, XL, Xl, xL, xl, Xm, XM, xm, xM, xN, xn, XN, Xn, xO, xo, Xo, XO, XP, Xp, xp, xP, XQ, Xq, xq, xQ, xr, xR, Xr, XR, xS, xs, Xs, XS, Xt, xt, XT, xT, Xu, XU, xu, xU, XV, xv, xV, Xv, xW, xw, XW, Xw, xx, xX, XX, Xx, xY, xy, Xy, XY, xZ, xz, Xz, XZ, y, Y, ya, yA, Ya, YA, YB, Yb, yb, yB, YC, Yc, yC, yc, yd, yD, Yd, YD, yE, ye, Ye, YE, yF, YF, Yf, yf, YG, Yg, yG, yg, yh, Yh, YH, yH, yI, yi, Yi, YI, YJ, yj, Yj, yJ, yk, yK, Yk, YK, yL, YL, Yl, yl, Ym, YM, yM, ym, YN, Yn, yn, yN, yO, yo, YO, Yo, yp, yP, YP, Yp, Yq, YQ, yQ, yq, Yr, YR, yR, yr, ys, yS, Ys, YS, yT, yt, YT, Yt, Yu, YU, yU, yu, Yv, YV, yV, yv, yw, yW, Yw, YW, yx, yX, YX, Yx, yY, Yy, yy, YY, yz, yZ, YZ, Yz, z, Z, zA, ZA, za, Za, zB, zb, ZB, Zb, Zc, ZC, zc, zC, Zd, ZD, zd, zD, ze, zE, ZE, Ze, zF, zf, Zf, ZF, Zg, ZG, zg, zG, ZH, Zh, zh, zH, zI, zi, Zi, ZI, zJ, zj, Zj, ZJ, Zk, zK, zk, ZK, zL, zl, Zl, ZL, zm, ZM, Zm, zM, ZN, Zn, zn, zN, Zo, ZO, zO, zo, zp, zP, ZP, Zp, zQ, zq, ZQ, Zq, ZR, Zr, zR, zr, Zs, ZS, zS, zs, zt, zT, Zt, ZT, zu, zU, Zu, ZU, ZV, Zv, zv, zV, ZW, Zw, zW, zw, zx, zX, Zx, ZX, zy, zY, Zy, ZY, ZZ, zz, zZ, Zz
            </forbidwords>
            <allowwords>i, pi, e, ket, depends</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options>allowEmpty, allowEmpty</options>
        </input>
        <input>
            <name>ans4</name>
            <type>algebraic</type>
            <tans>depends</tans>
            <boxsize>15</boxsize>
            <strictsyntax>1</strictsyntax>
            <insertstars>0</insertstars>
            <syntaxhint></syntaxhint>
            <syntaxattribute>0</syntaxattribute>
            <forbidwords>
                A, a, aA, Aa, AA, aa, ab, aB, Ab, AB, Ac, AC, aC, ac, aD, AD, Ad, ad, ae, AE, Ae, aE, Af, AF, af, aF, Ag, aG, AG, ag, aH, ah, Ah, AH, ai, aI, AI, Ai, AJ, Aj, aJ, aj, AK, Ak, aK, ak, al, aL, Al, AL, am, aM, Am, AM, AN, An, an, aN, AO, Ao, aO, ao, ap, aP, AP, Ap, aq, aQ, AQ, Aq, ar, aR, AR, Ar, aS, as, AS, As, At, AT, at, aT, Au, AU, au, aU, av, aV, Av, AV, aW, aw, Aw, AW, AX, ax, aX, Ax, AY, Ay, ay, aY, AZ, Az, aZ, az, b, B, ba, bA, BA, Ba, BB, Bb, bB, bb, bc, bC, BC, Bc, BD, Bd, bD, bd, Be, BE, bE, be, bF, bf, Bf, BF, BG, bg, bG, Bg, BH, Bh, bH, bh, bi, bI, BI, Bi, bj, bJ, Bj, BJ, Bk, BK, bk, bK, BL, Bl, bl, bL, bM, bm, BM, Bm, bN, bn, Bn, BN, Bo, BO, bO, bo, BP, Bp, bP, bp, bq, bQ, BQ, Bq, bR, br, Br, BR, bs, bS, Bs, BS, bt, bT, BT, Bt, Bu, BU, bU, bu, Bv, BV, bV, bv, bW, bw, BW, Bw, bX, bx, BX, Bx, by, bY, By, BY, Bz, BZ, bz, bZ, C, c, cA, ca, Ca, CA, cb, cB, CB, Cb, CC, Cc, cc, cC, cd, cD, CD, Cd, CE, Ce, ce, cE, CF, Cf, cF, cf, cg, cG, CG, Cg, Ch, cH, ch, CH, CI, Ci, cI, ci, cj, cJ, CJ, Cj, Ck, cK, ck, CK, Cl, CL, cL, cl, CM, cm, cM, Cm, cn, cN, Cn, CN, cO, co, Co, CO, CP, Cp, cp, cP, CQ, Cq, cq, cQ, cR, cr, CR, Cr, cS, cs, Cs, CS, ct, cT, CT, Ct, cu, cU, Cu, CU, CV, Cv, cv, cV, CW, Cw, cW, cw, cx, cX, Cx, CX, cy, cY, CY, Cy, Cz, cz, cZ, CZ, D, d, DA, Da, da, dA, dB, db, DB, Db, dC, dc, Dc, DC, Dd, DD, dD, dd, dE, de, De, DE, Df, DF, dF, df, DG, Dg, dG, dg, dH, dh, DH, Dh, DI, Di, di, dI, Dj, DJ, dj, dJ, dK, dk, Dk, DK, DL, Dl, dL, dl, Dm, DM, dm, dM, dn, dN, Dn, DN, do, dO, Do, DO, dP, DP, Dp, dp, Dq, DQ, dq, dQ, Dr, DR, dR, dr, ds, dS, Ds, DS, dt, dT, DT, Dt, dU, du, Du, DU, dV, dv, DV, Dv, DW, Dw, dW, dw, Dx, DX, dX, dx, dy, dY, Dy, DY, dz, dZ, DZ, Dz, E, Ea, EA, ea, eA, Eb, EB, eb, eB, ec, eC, EC, Ec, ed, eD, ED, Ed, Ee, EE, ee, eE, ef, eF, EF, Ef, EG, Eg, eg, eG, EH, Eh, eh, eH, ei, eI, EI, Ei, EJ, Ej, ej, eJ, ek, EK, Ek, eK, el, eL, El, EL, Em, EM, em, eM, EN, En, en, eN, Eo, eO, eo, EO, ep, eP, EP, Ep, EQ, Eq, eq, eQ, ER, Er, eR, er, ES, Es, eS, es, et, eT, Et, ET, eu, eU, Eu, EU, eV, ev, EV, Ev, ew, eW, EW, Ew, EX, Ex, ex, eX, EY, Ey, eY, ey, ez, eZ, EZ, Ez, f, F, FA, Fa, fa, fA, FB, Fb, fb, fB, Fc, FC, fC, fc, fD, fd, Fd, FD, fe, fE, FE, Fe, FF, Ff, fF, ff, fg, fG, Fg, FG, Fh, FH, fH, fh, FI, Fi, fI, fi, fJ, fj, FJ, Fj, Fk, fk, FK, fK, fL, fl, FL, Fl, fM, fm, FM, Fm, Fn, FN, fN, fn, Fo, FO, fO, fo, FP, Fp, fP, fp, fq, fQ, Fq, FQ, fr, Fr, FR, fR, Fs, FS, fs, fS, FT, Ft, ft, fT, fU, fu, Fu, FU, fV, fv, Fv, FV, fW, fw, FW, Fw, fX, fx, Fx, FX, FY, Fy, fy, fY, Fz, FZ, fz, fZ, G, g, gA, ga, Ga, GA, Gb, GB, gb, gB, Gc, GC, gc, gC, GD, Gd, gD, gd, ge, gE, Ge, GE, gF, gf, Gf, GF, Gg, GG, gG, gg, gH, gh, Gh, GH, Gi, GI, gI, gi, Gj, GJ, gJ, gj, gK, gk, Gk, GK, Gl, gl, GL, gL, gm, gM, GM, Gm, gn, gN, Gn, GN, GO, Go, gO, go, GP, Gp, gp, gP, GQ, gq, gQ, Gq, gR, gr, GR, Gr, Gs, gS, gs, GS, Gt, GT, gt, gT, GU, gu, Gu, gU, gv, gV, Gv, GV, gW, gw, GW, Gw, gx, gX, GX, Gx, gY, gy, Gy, GY, GZ, Gz, gZ, gz, H, h, ha, hA, Ha, HA, hb, hB, HB, Hb, Hc, HC, hC, hc, Hd, HD, hD, hd, HE, he, hE, He, hf, hF, HF, Hf, hG, Hg, hg, HG, Hh, HH, hh, hH, hI, Hi, hi, HI, HJ, Hj, hj, hJ, Hk, HK, hk, hK, hl, hL, HL, Hl, Hm, HM, hM, hm, hN, hn, HN, Hn, ho, hO, HO, Ho, Hp, HP, hP, hp, Hq, HQ, hQ, hq, hr, hR, Hr, HR, hs, hS, Hs, HS, HT, Ht, hT, ht, Hu, HU, hU, hu, hv, Hv, hV, HV, hW, hw, HW, Hw, hX, hx, HX, Hx, hy, hY, Hy, HY, hZ, HZ, Hz, hz, I, Ia, IA, ia, iA, ib, iB, Ib, IB, iC, ic, IC, Ic, ID, Id, id, iD, Ie, IE, ie, iE, if, If, iF, IF, iG, ig, IG, Ig, IH, Ih, iH, ih, II, Ii, iI, ii, ij, Ij, IJ, iJ, Ik, IK, iK, ik, IL, Il, il, iL, iM, im, IM, Im, IN, in, In, iN, iO, io, IO, Io, ip, iP, IP, Ip, IQ, Iq, iQ, iq, IR, Ir, ir, iR, iS, is, IS, Is, it, iT, It, IT, Iu, IU, iu, iU, Iv, IV, iv, iV, iW, iw, Iw, IW, ix, iX, Ix, IX, iy, iY, IY, Iy, iZ, iz, Iz, IZ, j, J, JA, Ja, ja, jA, JB, Jb, jB, jb, jC, jc, Jc, JC, jd, jD, Jd, JD, JE, Je, je, jE, JF, Jf, jF, jf, jG, jg, JG, Jg, jh, jH, Jh, JH, JI, ji, jI, Ji, Jj, JJ, jJ, jj, jk, jK, Jk, JK, JL, Jl, jL, jl, Jm, JM, jM, jm, jN, jn, Jn, JN, jo, jO, Jo, JO, jP, jp, Jp, JP, jq, jQ, JQ, Jq, Jr, JR, jr, jR, Js, JS, js, jS, jt, jT, JT, Jt, jU, ju, Ju, JU, Jv, JV, jV, jv, JW, Jw, jW, jw, jX, jx, JX, Jx, jy, jY, JY, Jy, jz, jZ, Jz, JZ, K, k, kA, ka, Ka, KA, Kb, KB, kb, kB, Kc, KC, kc, kC, kd, kD, KD, Kd, ke, kE, Ke, KE, Kf, KF, kF, kf, Kg, KG, kG, kg, kH, kh, KH, Kh, kI, ki, KI, Ki, kj, kJ, Kj, KJ, Kk, KK, kk, kK, kl, kL, Kl, KL, Km, KM, km, kM, KN, kn, kN, Kn, ko, kO, KO, Ko, kP, kp, KP, Kp, kQ, KQ, kq, Kq, kR, kr, Kr, KR, Ks, KS, kS, ks, Kt, KT, kt, kT, ku, kU, Ku, KU, kv, kV, KV, Kv, KW, Kw, kW, kw, Kx, KX, kx, kX, ky, kY, KY, Ky, kz, kZ, KZ, Kz, L, l, la, lA, LA, La, lb, lB, LB, Lb, Lc, LC, lC, lc, Ld, LD, lD, ld, le, lE, Le, LE, lf, lF, LF, Lf, LG, Lg, lG, lg, LH, Lh, lh, lH, li, lI, Li, LI, lj, lJ, Lj, LJ, Lk, lK, lk, LK, LL, Ll, ll, lL, lM, LM, Lm, lm, LN, Ln, ln, lN, lo, lO, LO, Lo, lp, lP, Lp, LP, lQ, Lq, LQ, lq, lr, lR, LR, Lr, ls, lS, LS, Ls, Lt, LT, lt, lT, Lu, LU, lU, lu, lV, lv, LV, Lv, lW, lw, LW, Lw, Lx, LX, lX, lx, LY, Ly, ly, lY, lz, lZ, Lz, LZ, m, M, MA, ma, mA, Ma, mb, mB, Mb, MB, mC, MC, Mc, mc, MD, Md, mD, md, Me, ME, me, mE, mF, mf, Mf, MF, mG, mg, MG, Mg, MH, Mh, mH, mh, Mi, MI, mI, mi, mj, mJ, Mj, MJ, mk, mK, MK, Mk, ML, Ml, ml, mL, Mm, MM, mM, mm, MN, Mn, mN, mn, MO, Mo, mO, mo, Mp, mP, mp, MP, mQ, mq, MQ, Mq, Mr, MR, mr, mR, mS, Ms, ms, MS, mt, mT, MT, Mt, MU, Mu, mU, mu, MV, Mv, mV, mv, mw, mW, MW, Mw, mx, mX, MX, Mx, MY, My, mY, my, MZ, Mz, mz, mZ, n, N, NA, Na, na, nA, Nb, nB, nb, NB, nC, nc, Nc, NC, nd, Nd, ND, nD, Ne, NE, ne, nE, NF, Nf, nF, nf, nG, ng, NG, Ng, nh, nH, NH, Nh, NI, Ni, ni, nI, NJ, Nj, nJ, nj, nK, nk, Nk, NK, nl, nL, Nl, NL, NM, Nm, nm, nM, NN, Nn, nN, nn, NO, No, no, nO, Np, NP, nP, np, NQ, Nq, nQ, nq, nr, nR, NR, Nr, NS, Ns, ns, nS, Nt, NT, nT, nt, nU, nu, Nu, NU, NV, Nv, nV, nv, Nw, NW, nW, nw, nX, nx, Nx, NX, ny, nY, NY, Ny, Nz, NZ, nz, nZ, O, o, Oa, OA, oA, oa, Ob, OB, ob, oB, Oc, OC, oc, oC, oD, od, OD, Od, oe, oE, Oe, OE, Of, OF, oF, of, Og, OG, oG, og, oh, oH, OH, Oh, oI, oi, Oi, OI, Oj, OJ, oj, oJ, Ok, OK, ok, oK, ol, oL, OL, Ol, om, oM, Om, OM, On, ON, oN, on, Oo, OO, oO, oo, Op, OP, oP, op, Oq, OQ, oQ, oq, OR, or, oR, Or, oS, os, Os, OS, ot, OT, oT, Ot, Ou, OU, ou, oU, oV, ov, OV, Ov, OW, Ow, ow, oW, OX, Ox, ox, oX, oY, oy, Oy, OY, Oz, oZ, oz, OZ, P, p, pa, pA, Pa, PA, PB, Pb, pb, pB, PC, pC, pc, Pc, Pd, PD, pD, pd, pe, pE, Pe, PE, pf, pF, PF, Pf, PG, pG, Pg, pg, PH, ph, pH, Ph, pI, Pi, PI, Pj, pj, pJ, PJ, PK, pK, pk, Pk, PL, Pl, pL, pl, pm, pM, Pm, PM, Pn, pn, pN, PN, PO, pO, po, Po, PP, Pp, pp, pP, PQ, Pq, pq, pQ, PR, pr, pR, Pr, Ps, pS, ps, PS, pt, Pt, pT, PT, pU, PU, Pu, pu, Pv, PV, pv, pV, pW, pw, PW, Pw, px, PX, pX, Px, Py, PY, py, pY, pZ, Pz, PZ, pz, q, Q, qa, qA, Qa, QA, qB, Qb, QB, qb, qc, Qc, QC, qC, qd, qD, QD, Qd, Qe, QE, qe, qE, Qf, qF, QF, qf, qG, qg, Qg, QG, qh, QH, Qh, qH, Qi, qI, qi, QI, qJ, Qj, QJ, qj, Qk, QK, qk, qK, ql, qL, QL, Ql, qm, Qm, QM, qM, qN, Qn, QN, qn, Qo, QO, qO, qo, qp, qP, QP, Qp, qq, Qq, QQ, qQ, qr, Qr, QR, qR, qs, qS, Qs, QS, QT, Qt, qT, qt, qU, QU, Qu, qu, qv, Qv, QV, qV, Qw, qw, QW, qW, qX, Qx, QX, qx, qy, qY, Qy, QY, QZ, Qz, qz, qZ, r, R, RA, Ra, rA, ra, rB, rb, RB, Rb, Rc, RC, rC, rc, rd, rD, Rd, RD, re, rE, RE, Re, RF, Rf, rf, rF, RG, Rg, rG, rg, rh, rH, Rh, RH, ri, rI, Ri, RI, RJ, rj, rJ, Rj, Rk, RK, rK, rk, Rl, RL, rL, rl, rm, rM, RM, Rm, rn, rN, RN, Rn, RO, Ro, rO, ro, RP, Rp, rp, rP, rq, rQ, RQ, Rq, rr, rR, RR, Rr, rs, rS, Rs, RS, rT, rt, RT, Rt, Ru, RU, ru, rU, Rv, RV, rv, rV, rw, rW, RW, Rw, RX, rX, rx, Rx, Ry, RY, ry, rY, rz, rZ, RZ, Rz, s, S, Sa, SA, sA, sa, Sb, SB, sb, sB, sc, sC, Sc, SC, SD, Sd, sd, sD, se, sE, Se, SE, sF, sf, Sf, SF, SG, Sg, sG, sg, SH, Sh, sH, sh, sI, si, SI, Si, sj, sJ, Sj, SJ, sk, sK, SK, Sk, SL, Sl, sL, sl, Sm, SM, sm, sM, sN, sn, Sn, SN, sO, so, SO, So, SP, Sp, sp, sP, Sq, SQ, sQ, sq, sr, sR, Sr, SR, ss, sS, SS, Ss, st, sT, ST, St, su, sU, SU, Su, SV, Sv, sV, sv, SW, Sw, sW, sw, sX, sx, Sx, SX, Sy, sY, SY, sy, SZ, Sz, sz, sZ, T, t, tA, ta, TA, Ta, TB, Tb, tB, tb, TC, Tc, tc, tC, tD, td, Td, TD, TE, Te, te, tE, tF, tf, TF, Tf, tG, tg, TG, Tg, TH, Th, th, tH, Ti, TI, ti, tI, tJ, tj, Tj, TJ, tK, tk, TK, Tk, Tl, tl, tL, TL, Tm, TM, tm, tM, TN, tN, tn, Tn, tO, to, To, TO, tp, tP, Tp, TP, TQ, Tq, tq, tQ, TR, Tr, tR, tr, tS, ts, Ts, TS, tT, tt, Tt, TT, tU, tu, Tu, TU, tV, tv, Tv, TV, Tw, TW, tW, tw, TX, Tx, tX, tx, ty, tY, Ty, TY, tz, TZ, Tz, tZ, u, U, ua, uA, Ua, UA, uB, Ub, UB, ub, UC, Uc, uc, uC, Ud, UD, ud, uD, ue, uE, Ue, UE, Uf, UF, uF, uf, uG, ug, Ug, UG, uh, uH, UH, Uh, Ui, UI, uI, ui, Uj, UJ, uj, uJ, uk, uK, Uk, UK, ul, uL, UL, Ul, Um, UM, uM, um, Un, UN, un, uN, uo, uO, Uo, UO, up, uP, UP, Up, uQ, Uq, UQ, uq, Ur, UR, ur, uR, Us, US, us, uS, uT, ut, UT, Ut, uu, uU, Uu, UU, uv, uV, Uv, UV, uw, uW, UW, Uw, UX, Ux, ux, uX, UY, Uy, uy, uY, uZ, uz, Uz, UZ, v, V, Va, vA, va, VA, vB, vb, VB, Vb, VC, Vc, vc, vC, Vd, VD, vd, vD, VE, Ve, vE, ve, vF, vf, VF, Vf, VG, Vg, vG, vg, vh, vH, VH, Vh, vi, vI, VI, Vi, VJ, Vj, vj, vJ, Vk, VK, vK, vk, vl, vL, Vl, VL, vm, vM, VM, Vm, VN, Vn, vn, vN, VO, Vo, vO, vo, Vp, vp, vP, VP, vq, vQ, VQ, Vq, VR, Vr, vr, vR, Vs, VS, vS, vs, Vt, VT, vT, vt, vu, vU, VU, Vu, vv, vV, VV, Vv, vW, vw, VW, Vw, vx, vX, Vx, VX, Vy, VY, vy, vY, VZ, Vz, vZ, vz, W, w, Wa, WA, wA, wa, WB, wb, wB, Wb, wc, wC, WC, Wc, WD, Wd, wD, wd, We, WE, wE, we, WF, Wf, wF, wf, wG, wg, WG, Wg, WH, Wh, wh, wH, wI, wi, Wi, WI, wj, wJ, Wj, WJ, WK, Wk, wk, wK, WL, Wl, wl, wL, wm, wM, Wm, WM, wN, wn, Wn, WN, Wo, WO, wO, wo, WP, Wp, wp, wP, WQ, wQ, wq, Wq, wr, wR, Wr, WR, Ws, WS, ws, wS, WT, Wt, wT, wt, Wu, WU, wu, wU, wV, wv, Wv, WV, wW, ww, Ww, WW, wX, wx, Wx, WX, wY, wy, Wy, WY, WZ, Wz, wz, wZ, X, x, Xa, XA, xA, xa, XB, Xb, xB, xb, xC, xc, Xc, XC, xd, xD, Xd, XD, xe, Xe, XE, xE, XF, Xf, xF, xf, XG, Xg, xG, xg, xh, xH, XH, Xh, Xi, XI, xi, xI, xJ, xj, XJ, Xj, xK, xk, Xk, XK, XL, Xl, xL, xl, Xm, XM, xm, xM, xN, xn, XN, Xn, xO, xo, Xo, XO, XP, Xp, xp, xP, XQ, Xq, xq, xQ, xr, xR, Xr, XR, xS, xs, Xs, XS, Xt, xt, XT, xT, Xu, XU, xu, xU, XV, xv, xV, Xv, xW, xw, XW, Xw, xx, xX, XX, Xx, xY, xy, Xy, XY, xZ, xz, Xz, XZ, y, Y, ya, yA, Ya, YA, YB, Yb, yb, yB, YC, Yc, yC, yc, yd, yD, Yd, YD, yE, ye, Ye, YE, yF, YF, Yf, yf, YG, Yg, yG, yg, yh, Yh, YH, yH, yI, yi, Yi, YI, YJ, yj, Yj, yJ, yk, yK, Yk, YK, yL, YL, Yl, yl, Ym, YM, yM, ym, YN, Yn, yn, yN, yO, yo, YO, Yo, yp, yP, YP, Yp, Yq, YQ, yQ, yq, Yr, YR, yR, yr, ys, yS, Ys, YS, yT, yt, YT, Yt, Yu, YU, yU, yu, Yv, YV, yV, yv, yw, yW, Yw, YW, yx, yX, YX, Yx, yY, Yy, yy, YY, yz, yZ, YZ, Yz, z, Z, zA, ZA, za, Za, zB, zb, ZB, Zb, Zc, ZC, zc, zC, Zd, ZD, zd, zD, ze, zE, ZE, Ze, zF, zf, Zf, ZF, Zg, ZG, zg, zG, ZH, Zh, zh, zH, zI, zi, Zi, ZI, zJ, zj, Zj, ZJ, Zk, zK, zk, ZK, zL, zl, Zl, ZL, zm, ZM, Zm, zM, ZN, Zn, zn, zN, Zo, ZO, zO, zo, zp, zP, ZP, Zp, zQ, zq, ZQ, Zq, ZR, Zr, zR, zr, Zs, ZS, zS, zs, zt, zT, Zt, ZT, zu, zU, Zu, ZU, ZV, Zv, zv, zV, ZW, Zw, zW, zw, zx, zX, Zx, ZX, zy, zY, Zy, ZY, ZZ, zz, zZ, Zz
            </forbidwords>
            <allowwords>i, pi, e, ket, depends</allowwords>
            <forbidfloat>1</forbidfloat>
            <requirelowestterms>0</requirelowestterms>
            <checkanswertype>0</checkanswertype>
            <mustverify>1</mustverify>
            <showvalidation>1</showvalidation>
            <options>allowEmpty, allowEmpty</options>
        </input>
    </question>
</quiz>